Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Выполните все задания и "запуште" решение

## Задание 1

Завершите функцию возвращающую в качестве результата `a`.

<!-- grade_id: cell-87f9daf4fba43d18, points: 0, grade: false, solution: true -->
```python
def fun1(a):
    return a
```

<!-- grade_id: cell-7ac55cd7df9bac79, points: 1, grade: true, solution: false -->
```python
'''test 1'''
assert fun1(1) == 1

'''test 2'''
assert fun1(2) == 2

```